Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Python for Data Science Bootcamp
## Assignment 1

NOTE: This assignment has 4 problems.  Please write your code in the designated portion of the autograded answer cell.

For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](https://github.com/Python-for-Data-Science-Bootcamp/a1-simrangulrajani4/blob/master/jupyter.png?raw=1)

©2020 Luis Morales. All rights reserved. 

## Problem 1 (1 point)

Write a function that test whether the input number is a positive odd integer. Don't worry about the incorrect input.

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
def myfunction(n):
    """test whether the input number is a positive odd integer. Don't worry about the incorrect input.
    """
    # YOUR CODE HERE
    if n>0 and n % 2 != 0:
      return True
    else: 
      return False

    raise NotImplementedError()
```

<!-- grade_id: evalmyfunction, points: 1, grade: true, solution: false -->
```python
"""
this is testing!
Check that positive_odd_int function returns the correct output for several inputs"""
assert myfunction(10) == False
assert myfunction(11) == True
assert myfunction(-21) == False
```

## Problem 2 (1 point)

Write a function that returns a list of numbers, such that $x_i=i^3$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: cubics, points: 0, grade: false, solution: true -->
```python
def cubics(n):
    """Compute the cubics of numbers from 1 to n, such that the 
    ith element of the returned list equals i^3.
    
    """
    # YOUR CODE HERE
    if n<1: 
        raise ValueError('Not a valid number')
    cubes=[]
    
    for i in range(1,n+1):
        cubes +=[i*i*i]
  
    return cubes
    raise NotImplementedError()
```

Your function should print `[1, 8, 27, 64, 125, 216, 343, 512, 729, 1000]` for $n=10$. Check that it does:

```python
cubics(10)
```

Output:
```
[1, 8, 27, 64, 125, 216, 343, 512, 729, 1000]
```

<!-- grade_id: correct_cubics, points: 0.75, grade: true, solution: false -->
```python
"""Check that cubics returns the correct output for several inputs"""
assert cubics(1) == [1]
assert cubics(2) == [1, 8]
assert cubics(10) == [1, 8, 27, 64, 125, 216, 343, 512, 729, 1000]
assert cubics(11) == [1, 8, 27, 64, 125, 216, 343, 512, 729, 1000, 1331]
```

<!-- grade_id: cubics_invalid_input, points: 0.25, grade: true, solution: false -->
```python
"""Check that cubics raises an error for invalid inputs"""
try:
    cubics(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    cubics(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

## Problem 3 (1 point)

Using your `cubics` function, write a function that computes the sum of the cubics of the numbers from 1 to $n$. Your function should call the `cubics` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_cubics, points: 0, grade: false, solution: true -->
```python
def sum_of_cubics(n):
    """Compute the sum of the cubics of numbers from 1 to n."""
    # YOUR CODE HERE
    cube_list=cubics(n)
    sum=0
    for i in range(len(cube_list)):
      sum+=cube_list[i]
    return sum
    raise NotImplementedError()
```

```python
sum_of_cubics(5)
```

Output:
```
225
```

<!-- grade_id: correct_sum_of_cubics, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_cubics returns the correct answer for various inputs."""
assert sum_of_cubics(1) == 1
assert sum_of_cubics(2) == 9
assert sum_of_cubics(10) == 3025
assert sum_of_cubics(11) == 4356
```

<!-- grade_id: sum_of_cubics_uses_cubics, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_cubics relies on cubics."""
orig_cubics = cubics
del cubics
try:
    sum_of_cubics(1)
except NameError:
    pass
else:
    raise AssertionError("sum_of_cubics does not use cubics")
finally:
    cubics = orig_cubics
```

## Problem 4 (1 point)

Write a function that returns a list of numbers, that form a Collatz conjecture sequence. Make sure the input is an int and it handles the case where $n<=0$ by raising a `ValueError`.
https://en.wikipedia.org/wiki/Collatz_conjecture

<!-- grade_id: Collatz, points: 0, grade: false, solution: true -->
```python
def Collatz(n):
    """For input n (must be a positive integer), compute a list of numbers in reverse order, that form a Collatz conjecture sequence.
  
    """
    # YOUR CODE HERE
    collat=[n]
    if not type(n)==int or n <= 0:
        raise ValueError('Invalid Number')
    while n > 1:
        if n % 2 == 0:
            n = n//2
        else:
            n = (3*n)+1
        collat += [n]
    return collat
    raise NotImplementedError()
```

```python

```

<!-- grade_id: Correct_Collatz, points: 0.5, grade: true, solution: false -->
```python
"""Check that cubics returns the correct output for several inputs"""
assert Collatz(1) == [1]
assert Collatz(2) == [2, 1]
assert Collatz(10) == [10, 5, 16, 8, 4, 2, 1]
assert Collatz(11) == [11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1]
```

<!-- grade_id: Collatz_invalid_input, points: 0.5, grade: true, solution: false -->
```python
"""Check that cubics raises an error for invalid inputs"""
try:
    Collatz(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    Collatz(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
    
try:
    Collatz(10.0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```